_64_add=off:afr=on:afp=100000:afq=1.1:anc=none:cond=fast:fde=none:irw=on:nm=6:newcnf=on:nwc=1.3:uhcvi=on_5 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 105414
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% lrs+10_6_aac=none:acc=model:add=large:afr=on:afp=10000:afq=1.2:amm=off:anc=none:bd=off:ccuc=small_ones:irw=on:lcm=reverse:nm=0:nwc=1:nicw=on:stl=30:sos=on:sp=reverse_arity:updr=off_2 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6396
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_1 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6140
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6140
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
% lrs+10_1_afr=on:afp=100000:afq=1.2:amm=sco:anc=none:br=off:cond=on:gs=on:gsem=on:irw=on:nm=16:nwc=1:stl=30:sac=on:sp=occurrence:urr=on:updr=off_12 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 218375
% Time elapsed: 1.400 s
% ------------------------------
% ------------------------------
% ott-3_3_av=off:cond=fast:fde=none:lcm=reverse:nm=6:nwc=1:sd=5:ss=axioms:st=2.0:sos=on:sp=reverse_arity:updr=off:uhcvi=on_17 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 1663
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% lrs+10_2:3_afp=1000:afq=1.1:amm=sco:anc=none:er=known:lcm=reverse:lwlo=on:nm=64:newcnf=on:nwc=2:stl=30:sd=5:ss=axioms:sos=theory:sac=on:sp=occurrence_233 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6140
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
% dis+11_32_av=off:ep=RST:fsr=off:lwlo=on:nm=6:nwc=1.1:sd=5:ss=axioms:st=5.0:sp=reverse_arity:uhcvi=on_2 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 1663
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% ott+1_5_afp=40000:afq=1.0:anc=all:fde=none:gs=on:irw=on:lma=on:nm=32:nwc=2:sos=all:sac=on:sp=occurrence:urr=ec_only:uhcvi=on_125 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 11129
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% ott+1_8:1_av=off:bd=preordered:bsr=on:lma=on:nm=64:newcnf=on:nwc=1.2:uhcvi=on_205 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 940155
% Time elapsed: 20.700 s
% ------------------------------
% ------------------------------
% lrs+1011_2:1_av=off:irw=on:lwlo=on:nm=16:newcnf=on:nwc=2:stl=30:sd=4:ss=axioms:st=3.0:sp=occurrence_30 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 1663
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% dis+1002_8_add=large:afp=100000:afq=1.2:amm=off:bs=on:irw=on:nm=2:newcnf=on:nwc=1.1:sos=on:sp=reverse_arity:urr=ec_only:updr=off_259 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6652
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% lrs+1011_5:4_acc=on:add=large:afr=on:afp=10000:afq=2.0:amm=sco:anc=none:bsr=on:ccuc=first:cond=on:fde=unused:gs=on:gsaa=from_current:gsem=off:irw=on:nm=2:newcnf=on:nwc=1.2:stl=30:sos=on:sac=on:sp=reverse_arity:updr=off_5 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 11129
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% dis+11_10_av=off:lma=on:nm=64:nwc=5:sp=reverse_arity_3 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 97098
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% ott-11_12_aac=none:afp=100000:afq=1.2:amm=sco:bs=on:bce=on:cond=fast:er=known:gs=on:gsaa=from_current:gsem=off:irw=on:nm=4:nwc=2:sas=z3:sos=all:sp=occurrence:urr=ec_only:updr=off_253 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6652
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
% ott+10_2:3_add=large:afp=40000:afq=1.1:amm=off:anc=all_dependent:bd=preordered:bs=unit_only:cond=fast:er=filter:gs=on:gsaa=from_current:lma=on:nm=32:nwc=1.1:sas=z3:sac=on:sp=occurrence:urr=ec_only:updr=off_679 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 128313
% Time elapsed: 68.100 s
% ------------------------------
% ------------------------------
% lrs+1011_3:2_aac=none:afp=10000:afq=1.2:amm=off:anc=all:cond=on:fde=none:gs=on:gsem=on:irw=on:lma=on:nm=32:newcnf=on:nwc=3:nicw=on:stl=30:sac=on:sp=reverse_arity:urr=ec_only:updr=off:uhcvi=on_21 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 240251
% Time elapsed: 2.300 s
% ------------------------------
% ------------------------------
% lrs+10_128_acc=model:afp=100000:afq=2.0:anc=all_dependent:bs=on:bsr=on:cond=fast:er=filter:gs=on:gsem=off:lcm=reverse:lma=on:nm=32:nwc=3:stl=30:sac=on:sp=occurrence:urr=ec_only_70 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 67674
% Time elapsed: 7.200 s
% ------------------------------
% ------------------------------
% lrs+1011_3:1_av=off:cond=on:irw=on:lma=on:nm=16:nwc=1:stl=30:sos=all:updr=off_51 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 2046
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% lrs+1010_5:4_av=off:bd=off:bsr=on:irw=on:lwlo=on:newcnf=on:nwc=1.1:stl=90:sos=all:sp=occurrence:uhcvi=on_145 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 2046
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% lrs+11_4_av=off:gsp=input_only:irw=on:lma=on:nm=0:nwc=1.2:stl=30:sd=2:ss=axioms:sp=reverse_arity:urr=on:updr=off_59 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 1663
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
% lrs+1011_14_av=off:fsr=off:irw=on:nwc=1:stl=30:sos=on:sp=occurrence:urr=ec_only:updr=off_53 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 2046
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% dis+11_5_afp=40000:afq=1.4:anc=none:br=off:cond=on:fsr=off:irw=on:lma=on:nm=4:nwc=1:sd=3:ss=axioms:st=2.0:urr=on:updr=off_5 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6140
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% lrs+1_3:2_aac=none:add=large:anc=all_dependent:bce=on:cond=fast:ep=RST:fsr=off:lma=on:nm=2:nwc=1:stl=30:sos=on:sp=occurrence:urr=on:updr=off:uhcvi=on_5 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 6652
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% lrs-11_8_afr=on:afp=100000:afq=2.0:anc=none:bd=off:cond=on:gs=on:lma=on:nm=2:newcnf=on:nwc=3:stl=30:sac=on:sp=occurrence_3 on 
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 43879
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% dis+1010_3:1_av=off:irw=on:nm=32:nwc=1:sos=all:urr=ec_only:updr=off_158 on 
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 2046
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% lrs+1010_3:2_afr=on:afp=100000:afq=1.1:anc=none:gsp=input_only:irw=on:lwlo=on:nm=2:newcnf=on:nwc=1.7:sac=on:sp=occurrence_300 on 
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for 
% SZS output start Proof for 
30. ! [X0] : term(q_Quote(X0)) [input]
36. ! [X0] : (term(X0) => term(q_Neg(X0))) [input]
39. ! [X0,X2] : ((term(X2) & term(X0)) => term(q_Forall(X0,X2))) [input]
40. term(c) [input]
41. term(q_c) [input]
42. term(q_X1) [input]
50. term(q_X4) [input]
51. term(q_X5) [input]
87. ! [X3] : (term(X3) => sub(q_X1,q_X1,X3) = X3) [input]
90. ! [X3] : (term(X3) => sub(q_X4,q_X4,X3) = X3) [input]
91. ! [X3] : (term(X3) => sub(q_X5,q_X5,X3) = X3) [input]
94. ! [X3] : (term(X3) => q_X1 = sub(q_X1,q_X4,X3)) [input]
95. ! [X3] : (term(X3) => q_X1 = sub(q_X1,q_X5,X3)) [input]
107. ! [X3] : (term(X3) => q_X4 = sub(q_X4,q_X5,X3)) [input]
109. ! [X3] : (term(X3) => q_X5 = sub(q_X5,q_X2,X3)) [input]
111. ! [X3] : (term(X3) => q_X5 = sub(q_X5,q_X4,X3)) [input]
132. ! [X0,X4] : (term(X0) => sub(q_Neg(X0),q_X1,X4) = q_Neg(sub(X0,q_X1,X4))) [input]
136. ! [X0,X4] : (term(X0) => sub(q_Neg(X0),q_X5,X4) = q_Neg(sub(X0,q_X5,X4))) [input]
147. ! [X0,X2,X4] : ((term(X2) & term(X0)) => sub(q_Forall(X0,X2),q_X1,X4) = q_Forall(sub(X0,q_X1,X4),sub(X2,q_X1,X4))) [input]
150. ! [X0,X2,X4] : ((term(X2) & term(X0)) => sub(q_Forall(X0,X2),q_X4,X4) = q_Forall(sub(X0,q_X4,X4),sub(X2,q_X4,X4))) [input]
151. ! [X0,X2,X4] : ((term(X2) & term(X0)) => sub(q_Forall(X0,X2),q_X5,X4) = q_Forall(sub(X0,q_X5,X4),sub(X2,q_X5,X4))) [input]
152. ! [X0,X2] : q_c = sub(q_c,X0,X2) [input]
153. ! [X0,X2] : ((term(X2) & term(X0)) => q_Forall(q_X1,X0) = sub(q_Forall(q_X1,X0),q_X1,X2)) [input]
157. ! [X0,X2] : ((term(X2) & term(X0)) => q_Forall(q_X5,X0) = sub(q_Forall(q_X5,X0),q_X5,X2)) [input]
165. ! [X0,X2] : ((term(X2) & term(X0)) => sub(q_Forall(q_X5,X0),q_X2,X2) = q_Forall(q_X5,sub(X0,q_X2,X2))) [input]
186. ! [X0] : (term(X0) => (q_Truth(q_Neg(X0)) <=> ~q_Truth(X0))) [input]
187. ! [X2] : ((term(X2) & term(X0)) => (q_Truth(q_Forall(q_X1,X2)) <=> ! [X3] : q_Truth(sub(X2,q_X1,q_Quote(X3))))) [input]
188. ! [X2] : ((term(X2) & term(X0)) => (q_Truth(q_Forall(q_X2,X2)) <=> ! [X3] : q_Truth(sub(X2,q_X2,q_Quote(X3))))) [input]
191. ! [X2] : ((term(X2) & term(X0)) => (q_Truth(q_Forall(q_X5,X2)) <=> ! [X3] : q_Truth(sub(X2,q_X5,q_Quote(X3))))) [input]
198. ! [X0] : (term(X0) => q_X1 = sub(q_X1,q_X4,X0)) [rectify 94]
200. ! [X0] : (term(X0) => q_X1 = sub(q_X1,q_X5,X0)) [rectify 95]
210. ! [X0] : (term(X0) => q_X5 = sub(q_X5,q_X2,X0)) [rectify 109]
212. ! [X0] : (term(X0) => q_X5 = sub(q_X5,q_X4,X0)) [rectify 111]
214. ! [X0] : (term(X0) => q_X4 = sub(q_X4,q_X5,X0)) [rectify 107]
218. ! [X0] : (term(X0) => sub(q_X5,q_X5,X0) = X0) [rectify 91]
219. ! [X0] : (term(X0) => sub(q_X4,q_X4,X0) = X0) [rectify 90]
221. ! [X0] : (term(X0) => sub(q_X1,q_X1,X0) = X0) [rectify 87]
223. ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X1,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X1,q_Quote(X2))))) [rectify 187]
224. ! [X1] : ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X1,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X1,q_Quote(X2))))) [closure 223]
225. ! [X1,X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X1,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X1,q_Quote(X2))))) [flattening 224]
226. ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X2,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X2,q_Quote(X2))))) [rectify 188]
227. ! [X1] : ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X2,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X2,q_Quote(X2))))) [closure 226]
228. ! [X1,X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X2,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X2,q_Quote(X2))))) [flattening 227]
229. ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X5,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X5,q_Quote(X2))))) [rectify 191]
230. ! [X1] : ! [X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X5,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X5,q_Quote(X2))))) [closure 229]
231. ! [X1,X0] : ((term(X0) & term(X1)) => (q_Truth(q_Forall(q_X5,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X5,q_Quote(X2))))) [flattening 230]
238. ! [X0,X1] : q_c = sub(q_c,X0,X1) [rectify 152]
256. ! [X0,X1] : (term(X0) => sub(q_Neg(X0),q_X1,X1) = q_Neg(sub(X0,q_X1,X1))) [rectify 132]
260. ! [X0,X1] : (term(X0) => sub(q_Neg(X0),q_X5,X1) = q_Neg(sub(X0,q_X5,X1))) [rectify 136]
281. ! [X0,X1] : ((term(X1) & term(X0)) => term(q_Forall(X0,X1))) [rectify 39]
288. ! [X0,X1] : ((term(X1) & term(X0)) => q_Forall(q_X5,X0) = sub(q_Forall(q_X5,X0),q_X5,X1)) [rectify 157]
290. ! [X0,X1] : ((term(X1) & term(X0)) => q_Forall(q_X1,X0) = sub(q_Forall(q_X1,X0),q_X1,X1)) [rectify 153]
311. ! [X0,X1] : ((term(X1) & term(X0)) => sub(q_Forall(q_X5,X0),q_X2,X1) = q_Forall(q_X5,sub(X0,q_X2,X1))) [rectify 165]
331. ! [X0,X1,X2] : ((term(X1) & term(X0)) => sub(q_Forall(X0,X1),q_X5,X2) = q_Forall(sub(X0,q_X5,X2),sub(X1,q_X5,X2))) [rectify 151]
332. ! [X0,X1,X2] : ((term(X1) & term(X0)) => sub(q_Forall(X0,X1),q_X1,X2) = q_Forall(sub(X0,q_X1,X2),sub(X1,q_X1,X2))) [rectify 147]
335. ! [X0,X1,X2] : ((term(X1) & term(X0)) => sub(q_Forall(X0,X1),q_X4,X2) = q_Forall(sub(X0,q_X4,X2),sub(X1,q_X4,X2))) [rectify 150]
356. ! [X0] : (term(q_Neg(X0)) | ~term(X0)) [ennf transformation 36]
359. ! [X0] : (q_X1 = sub(q_X1,q_X4,X0) | ~term(X0)) [ennf transformation 198]
361. ! [X0] : (q_X1 = sub(q_X1,q_X5,X0) | ~term(X0)) [ennf transformation 200]
371. ! [X0] : (q_X5 = sub(q_X5,q_X2,X0) | ~term(X0)) [ennf transformation 210]
373. ! [X0] : (q_X5 = sub(q_X5,q_X4,X0) | ~term(X0)) [ennf transformation 212]
375. ! [X0] : (q_X4 = sub(q_X4,q_X5,X0) | ~term(X0)) [ennf transformation 214]
380. ! [X0] : (sub(q_X5,q_X5,X0) = X0 | ~term(X0)) [ennf transformation 218]
381. ! [X0] : (sub(q_X4,q_X4,X0) = X0 | ~term(X0)) [ennf transformation 219]
383. ! [X0] : (sub(q_X1,q_X1,X0) = X0 | ~term(X0)) [ennf transformation 221]
386. ! [X0] : ((q_Truth(q_Neg(X0)) <=> ~q_Truth(X0)) | ~term(X0)) [ennf transformation 186]
387. ! [X1,X0] : ((q_Truth(q_Forall(q_X1,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X1,q_Quote(X2)))) | (~term(X0) | ~term(X1))) [ennf transformation 225]
388. ! [X1,X0] : ((q_Truth(q_Forall(q_X1,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X1,q_Quote(X2)))) | ~term(X0) | ~term(X1)) [flattening 387]
389. ! [X1,X0] : ((q_Truth(q_Forall(q_X2,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X2,q_Quote(X2)))) | (~term(X0) | ~term(X1))) [ennf transformation 228]
390. ! [X1,X0] : ((q_Truth(q_Forall(q_X2,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X2,q_Quote(X2)))) | ~term(X0) | ~term(X1)) [flattening 389]
391. ! [X1,X0] : ((q_Truth(q_Forall(q_X5,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X5,q_Quote(X2)))) | (~term(X0) | ~term(X1))) [ennf transformation 231]
392. ! [X1,X0] : ((q_Truth(q_Forall(q_X5,X0)) <=> ! [X2] : q_Truth(sub(X0,q_X5,q_Quote(X2)))) | ~term(X0) | ~term(X1)) [flattening 391]
412. ! [X0,X1] : (sub(q_Neg(X0),q_X1,X1) = q_Neg(sub(X0,q_X1,X1)) | ~term(X0)) [ennf transformation 256]
416. ! [X0,X1] : (sub(q_Neg(X0),q_X5,X1) = q_Neg(sub(X0,q_X5,X1)) | ~term(X0)) [ennf transformation 260]
450. ! [X0,X1] : (term(q_Forall(X0,X1)) | (~term(X1) | ~term(X0))) [ennf transformation 281]
451. ! [X0,X1] : (term(q_Forall(X0,X1)) | ~term(X1) | ~term(X0)) [flattening 450]
464. ! [X0,X1] : (q_Forall(q_X5,X0) = sub(q_Forall(q_X5,X0),q_X5,X1) | (~term(X1) | ~term(X0))) [ennf transformation 288]
465. ! [X0,X1] : (q_Forall(q_X5,X0) = sub(q_Forall(q_X5,X0),q_X5,X1) | ~term(X1) | ~term(X0)) [flattening 464]
468. ! [X0,X1] : (q_Forall(q_X1,X0) = sub(q_Forall(q_X1,X0),q_X1,X1) | (~term(X1) | ~term(X0))) [ennf transformation 290]
469. ! [X0,X1] : (q_Forall(q_X1,X0) = sub(q_Forall(q_X1,X0),q_X1,X1) | ~term(X1) | ~term(X0)) [flattening 468]
510. ! [X0,X1] : (sub(q_Forall(q_X5,X0),q_X2,X1) = q_Forall(q_X5,sub(X0,q_X2,X1)) | (~term(X1) | ~term(X0))) [ennf transformation 311]
511. ! [X0,X1] : (sub(q_Forall(q_X5,X0),q_X2,X1) = q_Forall(q_X5,sub(X0,q_X2,X1)) | ~term(X1) | ~term(X0)) [flattening 510]
547. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X5,X2) = q_Forall(sub(X0,q_X5,X2),sub(X1,q_X5,X2)) | (~term(X1) | ~term(X0))) [ennf transformation 331]
548. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X5,X2) = q_Forall(sub(X0,q_X5,X2),sub(X1,q_X5,X2)) | ~term(X1) | ~term(X0)) [flattening 547]
549. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X1,X2) = q_Forall(sub(X0,q_X1,X2),sub(X1,q_X1,X2)) | (~term(X1) | ~term(X0))) [ennf transformation 332]
550. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X1,X2) = q_Forall(sub(X0,q_X1,X2),sub(X1,q_X1,X2)) | ~term(X1) | ~term(X0)) [flattening 549]
555. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X4,X2) = q_Forall(sub(X0,q_X4,X2),sub(X1,q_X4,X2)) | (~term(X1) | ~term(X0))) [ennf transformation 335]
556. ! [X0,X1,X2] : (sub(q_Forall(X0,X1),q_X4,X2) = q_Forall(sub(X0,q_X4,X2),sub(X1,q_X4,X2)) | ~term(X1) | ~term(X0)) [flattening 555]
583. term(q_c) [cnf transformation 41]
584. term(c) [cnf transformation 40]
585. term(q_X1) [cnf transformation 42]
586. term(q_X5) [cnf transformation 51]
587. term(q_X4) [cnf transformation 50]
610. term(q_Quote(X0)) [cnf transformation 30]
622. term(q_Neg(X0)) | ~term(X0) [cnf transformation 356]
625. q_X1 = sub(q_X1,q_X4,X0) | ~term(X0) [cnf transformation 359]
627. q_X1 = sub(q_X1,q_X5,X0) | ~term(X0) [cnf transformation 361]
637. q_X5 = sub(q_X5,q_X2,X0) | ~term(X0) [cnf transformation 371]
639. q_X5 = sub(q_X5,q_X4,X0) | ~term(X0) [cnf transformation 373]
641. q_X4 = sub(q_X4,q_X5,X0) | ~term(X0) [cnf transformation 375]
646. sub(q_X5,q_X5,X0) = X0 | ~term(X0) [cnf transformation 380]
647. sub(q_X4,q_X4,X0) = X0 | ~term(X0) [cnf transformation 381]
649. sub(q_X1,q_X1,X0) = X0 | ~term(X0) [cnf transformation 383]
652. q_Truth(q_Neg(X0)) | q_Truth(X0) | ~term(X0) [cnf transformation 386]
653. ~q_Truth(q_Neg(X0)) | ~q_Truth(X0) | ~term(X0) [cnf transformation 386]
654. ~term(X1) | ~term(X0) | ~q_Truth(sub(X0,q_X1,q_Quote(sK0(X0)))) | q_Truth(q_Forall(q_X1,X0)) [cnf transformation 388]
655. ~term(X1) | ~term(X0) | q_Truth(sub(X0,q_X1,q_Quote(X2))) | ~q_Truth(q_Forall(q_X1,X0)) [cnf transformation 388]
656. ~term(X1) | ~term(X0) | ~q_Truth(sub(X0,q_X2,q_Quote(sK1(X0)))) | q_Truth(q_Forall(q_X2,X0)) [cnf transformation 390]
657. ~term(X1) | ~term(X0) | q_Truth(sub(X0,q_X2,q_Quote(X2))) | ~q_Truth(q_Forall(q_X2,X0)) [cnf transformation 390]
658. ~term(X1) | ~term(X0) | ~q_Truth(sub(X0,q_X5,q_Quote(sK2(X0)))) | q_Truth(q_Forall(q_X5,X0)) [cnf transformation 392]
659. ~term(X1) | ~term(X0) | q_Truth(sub(X0,q_X5,q_Quote(X2))) | ~q_Truth(q_Forall(q_X5,X0)) [cnf transformation 392]
664. q_c = sub(q_c,X0,X1) [cnf transformation 238]
682. sub(q_Neg(X0),q_X1,X1) = q_Neg(sub(X0,q_X1,X1)) | ~term(X0) [cnf transformation 412]
686. sub(q_Neg(X0),q_X5,X1) = q_Neg(sub(X0,q_X5,X1)) | ~term(X0) [cnf transformation 416]
708. term(q_Forall(X0,X1)) | ~term(X1) | ~term(X0) [cnf transformation 451]
715. q_Forall(q_X5,X0) = sub(q_Forall(q_X5,X0),q_X5,X1) | ~term(X1) | ~term(X0) [cnf transformation 465]
717. q_Forall(q_X1,X0) = sub(q_Forall(q_X1,X0),q_X1,X1) | ~term(X1) | ~term(X0) [cnf transformation 469]
738. sub(q_Forall(q_X5,X0),q_X2,X1) = q_Forall(q_X5,sub(X0,q_X2,X1)) | ~term(X1) | ~term(X0) [cnf transformation 511]
762. sub(q_Forall(X0,X1),q_X5,X2) = q_Forall(sub(X0,q_X5,X2),sub(X1,q_X5,X2)) | ~term(X1) | ~term(X0) [cnf transformation 548]
763. sub(q_Forall(X0,X1),q_X1,X2) = q_Forall(sub(X0,q_X1,X2),sub(X1,q_X1,X2)) | ~term(X1) | ~term(X0) [cnf transformation 550]
766. sub(q_Forall(X0,X1),q_X4,X2) = q_Forall(sub(X0,q_X4,X2),sub(X1,q_X4,X2)) | ~term(X1) | ~term(X0) [cnf transformation 556]
802. ~term(X1) | sP5 [general splitting component introduction]
803. ~term(X0) | q_Truth(sub(X0,q_X1,q_Quote(X2))) | ~q_Truth(q_Forall(q_X1,X0)) | ~sP5 [general splitting 655,802]
804. ~term(X0) | ~q_Truth(sub(X0,q_X1,q_Quote(sK0(X0)))) | q_Truth(q_Forall(q_X1,X0)) | sP6 [general splitting component introduction]
805. ~term(X1) | ~sP6 [general splitting 654,804]
806. ~term(X1) | sP7 [general splitting component introduction]
807. ~term(X0) | q_Truth(sub(X0,q_X2,q_Quote(X2))) | ~q_Truth(q_Forall(q_X2,X0)) | ~sP7 [general splitting 657,806]
808. ~term(X0) | ~q_Truth(sub(X0,q_X2,q_Quote(sK1(X0)))) | q_Truth(q_Forall(q_X2,X0)) | sP8 [general splitting component introduction]
809. ~term(X1) | ~sP8 [general splitting 656,808]
810. ~term(X1) | sP9 [general splitting component introduction]
811. ~term(X0) | q_Truth(sub(X0,q_X5,q_Quote(X2))) | ~q_Truth(q_Forall(q_X5,X0)) | ~sP9 [general splitting 659,810]
812. ~term(X0) | ~q_Truth(sub(X0,q_X5,q_Quote(sK2(X0)))) | q_Truth(q_Forall(q_X5,X0)) | sP10 [general splitting component introduction]
813. ~term(X1) | ~sP10 [general splitting 658,812]
824. ~q_Truth(sub(X0,q_X5,q_Quote(sK2(X0)))) | ~term(X0) | q_Truth(q_Forall(q_X5,X0)) [subsumption resolution 812,813]
825. ~q_Truth(sub(X0,q_X2,q_Quote(sK1(X0)))) | ~term(X0) | q_Truth(q_Forall(q_X2,X0)) [subsumption resolution 808,809]
826. ~q_Truth(sub(X0,q_X1,q_Quote(sK0(X0)))) | ~term(X0) | q_Truth(q_Forall(q_X1,X0)) [subsumption resolution 804,805]
828. 1 <=> sP5 [avatar definition]
832. 2 <=> ! [X1] : ~term(X1) [avatar definition]
833. ~term(X1) <- (2) [avatar component clause 832]
834. 1 | 2 [avatar split clause 802,832,828]
835. $false <- (2) [resolution 833,584]
850. ~2 [avatar contradiction clause 835]
857. 4 <=> sP7 [avatar definition]
860. 4 | 2 [avatar split clause 806,832,857]
867. 6 <=> sP9 [avatar definition]
870. 6 | 2 [avatar split clause 810,832,867]
1092. 25 <=> ! [X0,X2] : (~term(X0) | ~q_Truth(q_Forall(q_X1,X0)) | q_Truth(sub(X0,q_X1,q_Quote(X2)))) [avatar definition]
1093. q_Truth(sub(X0,q_X1,q_Quote(X2))) | ~q_Truth(q_Forall(q_X1,X0)) | ~term(X0) <- (25) [avatar component clause 1092]
1094. ~1 | 25 [avatar split clause 803,1092,828]
1096. q_Truth(q_Neg(sub(X2,q_X1,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X1,q_Neg(X2))) | ~term(q_Neg(X2)) | ~term(X2) <- (25) [superposition 1093,682]
1107. q_Truth(q_Neg(sub(X2,q_X1,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X1,q_Neg(X2))) | ~term(X2) <- (25) [subsumption resolution 1096,622]
1124. 26 <=> q_Truth(q_Forall(q_X1,q_c)) [avatar definition]
1125. q_Truth(q_Forall(q_X1,q_c)) <- (26) [avatar component clause 1124]
1126. ~q_Truth(q_Forall(q_X1,q_c)) <- (~26) [avatar component clause 1124]
1133. 28 <=> ! [X0,X2] : (~term(X0) | ~q_Truth(q_Forall(q_X2,X0)) | q_Truth(sub(X0,q_X2,q_Quote(X2)))) [avatar definition]
1134. q_Truth(sub(X0,q_X2,q_Quote(X2))) | ~q_Truth(q_Forall(q_X2,X0)) | ~term(X0) <- (28) [avatar component clause 1133]
1135. ~4 | 28 [avatar split clause 807,1133,857]
1165. 29 <=> ! [X0,X2] : (~term(X0) | ~q_Truth(q_Forall(q_X5,X0)) | q_Truth(sub(X0,q_X5,q_Quote(X2)))) [avatar definition]
1166. q_Truth(sub(X0,q_X5,q_Quote(X2))) | ~q_Truth(q_Forall(q_X5,X0)) | ~term(X0) <- (29) [avatar component clause 1165]
1167. ~6 | 29 [avatar split clause 811,1165,867]
1178. q_Truth(q_Neg(sub(X2,q_X5,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X5,q_Neg(X2))) | ~term(q_Neg(X2)) | ~term(X2) <- (29) [superposition 1166,686]
1187. q_Truth(q_Quote(X15)) | ~q_Truth(q_Forall(q_X5,q_X5)) | ~term(q_X5) | ~term(q_Quote(X15)) <- (29) [superposition 1166,646]
1189. q_Truth(q_Neg(sub(X2,q_X5,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X5,q_Neg(X2))) | ~term(X2) <- (29) [subsumption resolution 1178,622]
1203. q_Truth(q_Quote(X15)) | ~q_Truth(q_Forall(q_X5,q_X5)) | ~term(q_Quote(X15)) <- (29) [subsumption resolution 1187,586]
1204. q_Truth(q_Quote(X15)) | ~q_Truth(q_Forall(q_X5,q_X5)) <- (29) [subsumption resolution 1203,610]
1349. ~q_Truth(q_Neg(sub(X1,q_X5,q_Quote(sK2(q_Neg(X1)))))) | ~term(q_Neg(X1)) | q_Truth(q_Forall(q_X5,q_Neg(X1))) | ~term(X1) [superposition 824,686]
1361. ~q_Truth(q_Neg(sub(X1,q_X5,q_Quote(sK2(q_Neg(X1)))))) | q_Truth(q_Forall(q_X5,q_Neg(X1))) | ~term(X1) [subsumption resolution 1349,622]
1418. ~q_Truth(q_Neg(sub(X1,q_X1,q_Quote(sK0(q_Neg(X1)))))) | ~term(q_Neg(X1)) | q_Truth(q_Forall(q_X1,q_Neg(X1))) | ~term(X1) [superposition 826,682]
1422. ~q_Truth(q_c) | ~term(q_c) | q_Truth(q_Forall(q_X1,q_c)) [superposition 826,664]
1430. ~q_Truth(q_Neg(sub(X1,q_X1,q_Quote(sK0(q_Neg(X1)))))) | q_Truth(q_Forall(q_X1,q_Neg(X1))) | ~term(X1) [subsumption resolution 1418,622]
1435. ~q_Truth(q_c) | q_Truth(q_Forall(q_X1,q_c)) [subsumption resolution 1422,583]
1436. ~q_Truth(q_c) <- (~26) [subsumption resolution 1435,1126]
1738. 47 <=> q_Truth(q_Forall(q_X5,q_X1)) [avatar definition]
1739. q_Truth(q_Forall(q_X5,q_X1)) <- (47) [avatar component clause 1738]
1740. ~q_Truth(q_Forall(q_X5,q_X1)) <- (~47) [avatar component clause 1738]
1778. q_Truth(q_Forall(q_X5,sub(X2,q_X2,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,X2))) | ~term(q_Forall(q_X5,X2)) | ~term(q_Quote(X3)) | ~term(X2) <- (28) [superposition 1134,738]
1779. ~q_Truth(q_Forall(q_X5,sub(X4,q_X2,q_Quote(sK1(q_Forall(q_X5,X4)))))) | ~term(q_Forall(q_X5,X4)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,X4))) | ~term(q_Quote(sK1(q_Forall(q_X5,X4)))) | ~term(X4) [superposition 825,738]
1781. q_Truth(q_Forall(q_X5,sub(X2,q_X2,q_Quote(X3)))) | ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,X2))) | ~term(q_Forall(q_X5,X2)) | ~term(X2) <- (28) [subsumption resolution 1778,610]
1782. ~q_Truth(q_Forall(q_X5,sub(X4,q_X2,q_Quote(sK1(q_Forall(q_X5,X4)))))) | ~term(q_Forall(q_X5,X4)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,X4))) | ~term(X4) [subsumption resolution 1779,610]
2749. sub(q_Forall(q_X5,X41),q_X5,X40) = q_Forall(X40,sub(X41,q_X5,X40)) | ~term(X41) | ~term(q_X5) | ~term(X40) [superposition 762,646]
2760. sub(q_Forall(X31,q_c),q_X5,X30) = q_Forall(sub(X31,q_X5,X30),q_c) | ~term(q_c) | ~term(X31) [superposition 762,664]
2761. sub(q_Forall(X33,q_X1),q_X5,X32) = q_Forall(sub(X33,q_X5,X32),q_X1) | ~term(q_X1) | ~term(X33) | ~term(X32) [superposition 762,627]
2764. sub(q_Forall(X39,q_X4),q_X5,X38) = q_Forall(sub(X39,q_X5,X38),q_X4) | ~term(q_X4) | ~term(X39) | ~term(X38) [superposition 762,641]
2765. sub(q_Forall(X41,q_X5),q_X5,X40) = q_Forall(sub(X41,q_X5,X40),X40) | ~term(q_X5) | ~term(X41) | ~term(X40) [superposition 762,646]
2775. sub(q_Forall(q_X5,X41),q_X5,X40) = q_Forall(X40,sub(X41,q_X5,X40)) | ~term(X41) | ~term(X40) [subsumption resolution 2749,586]
2781. sub(q_Forall(X31,q_c),q_X5,X30) = q_Forall(sub(X31,q_X5,X30),q_c) | ~term(X31) [subsumption resolution 2760,583]
2782. sub(q_Forall(X33,q_X1),q_X5,X32) = q_Forall(sub(X33,q_X5,X32),q_X1) | ~term(X33) | ~term(X32) [subsumption resolution 2761,585]
2785. sub(q_Forall(X39,q_X4),q_X5,X38) = q_Forall(sub(X39,q_X5,X38),q_X4) | ~term(X39) | ~term(X38) [subsumption resolution 2764,587]
2786. sub(q_Forall(X41,q_X5),q_X5,X40) = q_Forall(sub(X41,q_X5,X40),X40) | ~term(X41) | ~term(X40) [subsumption resolution 2765,586]
2798. sub(q_Forall(q_X1,X33),q_X1,X32) = q_Forall(X32,sub(X33,q_X1,X32)) | ~term(X33) | ~term(q_X1) | ~term(X32) [superposition 763,649]
2813. sub(q_Forall(X31,q_c),q_X1,X30) = q_Forall(sub(X31,q_X1,X30),q_c) | ~term(q_c) | ~term(X31) [superposition 763,664]
2814. sub(q_Forall(X33,q_X1),q_X1,X32) = q_Forall(sub(X33,q_X1,X32),X32) | ~term(q_X1) | ~term(X33) | ~term(X32) [superposition 763,649]
2828. sub(q_Forall(q_X1,X33),q_X1,X32) = q_Forall(X32,sub(X33,q_X1,X32)) | ~term(X33) | ~term(X32) [subsumption resolution 2798,585]
2834. sub(q_Forall(X31,q_c),q_X1,X30) = q_Forall(sub(X31,q_X1,X30),q_c) | ~term(X31) [subsumption resolution 2813,583]
2835. sub(q_Forall(X33,q_X1),q_X1,X32) = q_Forall(sub(X33,q_X1,X32),X32) | ~term(X33) | ~term(X32) [subsumption resolution 2814,585]
2961. sub(q_Forall(q_c,X31),q_X4,X30) = q_Forall(q_c,sub(X31,q_X4,X30)) | ~term(X31) | ~term(q_c) [superposition 766,664]
2981. sub(q_Forall(X39,q_X4),q_X4,X38) = q_Forall(sub(X39,q_X4,X38),X38) | ~term(q_X4) | ~term(X39) | ~term(X38) [superposition 766,647]
2991. sub(q_Forall(q_c,X31),q_X4,X30) = q_Forall(q_c,sub(X31,q_X4,X30)) | ~term(X31) [subsumption resolution 2961,583]
3002. sub(q_Forall(X39,q_X4),q_X4,X38) = q_Forall(sub(X39,q_X4,X38),X38) | ~term(X39) | ~term(X38) [subsumption resolution 2981,587]
3054. 62 <=> ! [X11] : q_Truth(q_Quote(X11)) [avatar definition]
3055. q_Truth(q_Quote(X11)) <- (62) [avatar component clause 3054]
3063. 64 <=> q_Truth(q_Forall(q_X5,q_X5)) [avatar definition]
3064. q_Truth(q_Forall(q_X5,q_X5)) <- (64) [avatar component clause 3063]
3065. ~q_Truth(q_Forall(q_X5,q_X5)) <- (~64) [avatar component clause 3063]
3066. ~64 | 62 | ~29 [avatar split clause 1204,1165,3054,3063]
3586. q_Truth(q_Neg(q_c)) | ~q_Truth(q_Forall(q_X1,q_Neg(q_c))) | ~term(q_c) <- (25) [superposition 1107,664]
3603. q_Truth(q_Neg(q_c)) | ~q_Truth(q_Forall(q_X1,q_Neg(q_c))) <- (25) [subsumption resolution 3586,583]
3917. q_Truth(q_Neg(q_Quote(X25))) | ~q_Truth(q_Forall(q_X5,q_Neg(q_X5))) | ~term(q_X5) | ~term(q_Quote(X25)) <- (29) [superposition 1189,646]
3938. q_Truth(q_Neg(q_Quote(X25))) | ~q_Truth(q_Forall(q_X5,q_Neg(q_X5))) | ~term(q_Quote(X25)) <- (29) [subsumption resolution 3917,586]
3939. q_Truth(q_Neg(q_Quote(X25))) | ~q_Truth(q_Forall(q_X5,q_Neg(q_X5))) <- (29) [subsumption resolution 3938,610]
4615. 132 <=> q_Truth(q_Forall(q_X1,q_Neg(q_c))) [avatar definition]
4617. ~q_Truth(q_Forall(q_X1,q_Neg(q_c))) <- (~132) [avatar component clause 4615]
4619. 133 <=> q_Truth(q_Neg(q_c)) [avatar definition]
4622. ~132 | 133 | ~25 [avatar split clause 3603,1092,4619,4615]
5134. q_Forall(q_X1,q_c) = q_Forall(sub(q_X1,q_X1,X3),q_c) | ~term(q_X1) | ~term(X3) | ~term(q_c) [superposition 2834,717]
5163. q_Forall(q_X1,q_c) = q_Forall(sub(q_X1,q_X1,X3),q_c) | ~term(X3) | ~term(q_c) [subsumption resolution 5134,585]
5164. q_Forall(q_X1,q_c) = q_Forall(sub(q_X1,q_X1,X3),q_c) | ~term(X3) [subsumption resolution 5163,583]
5898. ~q_Truth(q_Neg(q_c)) | q_Truth(q_Forall(q_X5,q_Neg(q_c))) | ~term(q_c) [superposition 1361,664]
5916. ~q_Truth(q_Neg(q_c)) | q_Truth(q_Forall(q_X5,q_Neg(q_c))) [subsumption resolution 5898,583]
6299. ~q_Truth(q_Neg(q_c)) | q_Truth(q_Forall(q_X1,q_Neg(q_c))) | ~term(q_c) [superposition 1430,664]
6317. ~q_Truth(q_Neg(q_c)) | ~term(q_c) <- (~132) [subsumption resolution 6299,4617]
6318. ~q_Truth(q_Neg(q_c)) <- (~132) [subsumption resolution 6317,583]
6443. q_Truth(q_c) | ~term(q_c) <- (~132) [resolution 6318,652]
6444. ~term(q_c) <- (~26, ~132) [subsumption resolution 6443,1436]
6445. $false <- (~26, ~132) [subsumption resolution 6444,583]
6446. 26 | 132 [avatar contradiction clause 6445]
11899. 185 <=> q_Truth(q_Forall(q_X5,q_Neg(q_c))) [avatar definition]
11900. q_Truth(q_Forall(q_X5,q_Neg(q_c))) <- (185) [avatar component clause 11899]
19541. q_Forall(sub(q_X5,q_X5,X4),q_c) = q_Forall(X4,sub(q_c,q_X5,X4)) | ~term(q_c) | ~term(X4) | ~term(q_X5) [superposition 2775,2781]
19592. sub(q_Forall(q_X5,q_X1),q_X5,X99) = q_Forall(X99,q_X1) | ~term(q_X1) | ~term(X99) | ~term(X99) [superposition 2775,627]
19596. sub(q_Forall(q_X5,q_X5),q_X5,X103) = q_Forall(X103,X103) | ~term(q_X5) | ~term(X103) | ~term(X103) [superposition 2775,646]
19625. sub(q_Forall(q_X5,q_X5),q_X5,X103) = q_Forall(X103,X103) | ~term(q_X5) | ~term(X103) [duplicate literal removal 19596]
19629. sub(q_Forall(q_X5,q_X1),q_X5,X99) = q_Forall(X99,q_X1) | ~term(q_X1) | ~term(X99) [duplicate literal removal 19592]
19668. q_Forall(sub(q_X5,q_X5,X4),q_c) = q_Forall(X4,sub(q_c,q_X5,X4)) | ~term(X4) | ~term(q_X5) [subsumption resolution 19541,583]
19669. q_Forall(sub(q_X5,q_X5,X4),q_c) = q_Forall(X4,sub(q_c,q_X5,X4)) | ~term(X4) [subsumption resolution 19668,586]
19670. q_Forall(sub(q_X5,q_X5,X4),q_c) = q_Forall(X4,q_c) | ~term(X4) [forward demodulation 19669,664]
19677. sub(q_Forall(q_X5,q_X1),q_X5,X99) = q_Forall(X99,q_X1) | ~term(X99) [subsumption resolution 19629,585]
19681. sub(q_Forall(q_X5,q_X5),q_X5,X103) = q_Forall(X103,X103) | ~term(X103) [subsumption resolution 19625,586]
19696. q_Forall(q_X5,q_X1) = q_Forall(sub(q_X5,q_X5,X8),q_X1) | ~term(q_X5) | ~term(X8) | ~term(X8) | ~term(q_X1) [superposition 2782,715]
19747. q_Forall(q_X5,q_X1) = q_Forall(sub(q_X5,q_X5,X8),q_X1) | ~term(q_X5) | ~term(X8) | ~term(q_X1) [duplicate literal removal 19696]
19765. q_Forall(q_X5,q_X1) = q_Forall(sub(q_X5,q_X5,X8),q_X1) | ~term(X8) | ~term(q_X1) [subsumption resolution 19747,586]
19766. q_Forall(q_X5,q_X1) = q_Forall(sub(q_X5,q_X5,X8),q_X1) | ~term(X8) [subsumption resolution 19765,585]
19986. q_Forall(q_X5,q_X4) = q_Forall(sub(q_X5,q_X5,X8),q_X4) | ~term(q_X5) | ~term(X8) | ~term(X8) | ~term(q_X4) [superposition 2785,715]
20037. q_Forall(q_X5,q_X4) = q_Forall(sub(q_X5,q_X5,X8),q_X4) | ~term(q_X5) | ~term(X8) | ~term(q_X4) [duplicate literal removal 19986]
20055. q_Forall(q_X5,q_X4) = q_Forall(sub(q_X5,q_X5,X8),q_X4) | ~term(X8) | ~term(q_X4) [subsumption resolution 20037,586]
20056. q_Forall(q_X5,q_X4) = q_Forall(sub(q_X5,q_X5,X8),q_X4) | ~term(X8) [subsumption resolution 20055,587]
20081. q_Forall(q_X5,q_X5) = q_Forall(sub(q_X5,q_X5,X8),X8) | ~term(q_X5) | ~term(X8) | ~term(X8) | ~term(q_X5) [superposition 2786,715]
20229. q_Forall(q_X5,q_X5) = q_Forall(sub(q_X5,q_X5,X8),X8) | ~term(q_X5) | ~term(X8) [duplicate literal removal 20081]
20247. q_Forall(q_X5,q_X5) = q_Forall(sub(q_X5,q_X5,X8),X8) | ~term(X8) [subsumption resolution 20229,586]
20262. q_Forall(sub(q_X1,q_X1,X4),q_c) = q_Forall(X4,sub(q_c,q_X1,X4)) | ~term(q_c) | ~term(X4) | ~term(q_X1) [superposition 2828,2834]
20389. q_Forall(sub(q_X1,q_X1,X4),q_c) = q_Forall(X4,sub(q_c,q_X1,X4)) | ~term(X4) | ~term(q_X1) [subsumption resolution 20262,583]
20390. q_Forall(sub(q_X1,q_X1,X4),q_c) = q_Forall(X4,sub(q_c,q_X1,X4)) | ~term(X4) [subsumption resolution 20389,585]
20391. q_Forall(X4,q_c) = q_Forall(sub(q_X1,q_X1,X4),q_c) | ~term(X4) [forward demodulation 20390,664]
20413. q_Forall(q_X1,q_X1) = q_Forall(sub(q_X1,q_X1,X4),X4) | ~term(q_X1) | ~term(X4) | ~term(X4) | ~term(q_X1) [superposition 2835,717]
20561. q_Forall(q_X1,q_X1) = q_Forall(sub(q_X1,q_X1,X4),X4) | ~term(q_X1) | ~term(X4) [duplicate literal removal 20413]
20567. q_Forall(q_X1,q_X1) = q_Forall(sub(q_X1,q_X1,X4),X4) | ~term(X4) [subsumption resolution 20561,585]
22929. sub(q_Forall(q_X5,q_X4),q_X4,X111) = q_Forall(q_X5,X111) | ~term(q_X5) | ~term(X111) | ~term(X111) [superposition 3002,639]
22970. sub(q_Forall(q_X5,q_X4),q_X4,X111) = q_Forall(q_X5,X111) | ~term(q_X5) | ~term(X111) [duplicate literal removal 22929]
23044. sub(q_Forall(q_X5,q_X4),q_X4,X111) = q_Forall(q_X5,X111) | ~term(X111) [subsumption resolution 22970,586]
26887. q_Truth(q_Forall(q_X5,q_X5)) | ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_Forall(q_X5,q_X5)) | ~term(q_X5) | ~term(q_Quote(X113)) <- (28) [superposition 1781,637]
26932. ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_Forall(q_X5,q_X5)) | ~term(q_X5) | ~term(q_Quote(X113)) <- (28, ~64) [subsumption resolution 26887,3065]
26933. ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_Forall(q_X5,q_X5)) | ~term(q_Quote(X113)) <- (28, ~64) [subsumption resolution 26932,586]
26934. ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_Forall(q_X5,q_X5)) <- (28, ~64) [subsumption resolution 26933,610]
32001. 185 | ~133 [avatar split clause 5916,4619,11899]
40854. 243 <=> ! [X21] : q_Truth(q_Neg(q_Quote(X21))) [avatar definition]
40855. q_Truth(q_Neg(q_Quote(X21))) <- (243) [avatar component clause 40854]
45074. 253 <=> q_Truth(q_Forall(q_X5,q_Neg(q_X5))) [avatar definition]
45076. ~q_Truth(q_Forall(q_X5,q_Neg(q_X5))) <- (~253) [avatar component clause 45074]
45077. ~253 | 243 | ~29 [avatar split clause 3939,1165,40854,45074]
65865. ~q_Truth(q_Forall(q_X5,q_X5)) | ~term(q_Forall(q_X5,q_X5)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_X5) | ~term(q_Quote(sK1(q_Forall(q_X5,q_X5)))) [superposition 1782,637]
100502. 286 <=> term(q_Forall(q_X5,q_X5)) [avatar definition]
100504. ~term(q_Forall(q_X5,q_X5)) <- (~286) [avatar component clause 100502]
101175. ~term(q_X5) | ~term(q_X5) <- (~286) [resolution 100504,708]
101176. ~term(q_X5) <- (~286) [duplicate literal removal 101175]
101177. $false <- (~286) [subsumption resolution 101176,586]
101178. 286 [avatar contradiction clause 101177]
136023. 320 <=> q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) [avatar definition]
136025. ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) <- (~320) [avatar component clause 136023]
136026. ~286 | ~320 | ~28 | 64 [avatar split clause 26934,3063,1133,136023,100502]
195817. q_Forall(q_X5,q_X1) = q_Forall(X2,q_X1) | ~term(X2) | ~term(X2) | ~term(q_X1) [superposition 19677,715]
195944. q_Forall(q_X5,q_X1) = q_Forall(X2,q_X1) | ~term(X2) | ~term(q_X1) [duplicate literal removal 195817]
195949. q_Forall(q_X5,q_X1) = q_Forall(X2,q_X1) | ~term(X2) [subsumption resolution 195944,585]
196208. term(q_Forall(q_X5,q_X1)) | ~term(q_X1) | ~term(X12) | ~term(X12) [superposition 708,195949]
196249. term(q_Forall(q_X5,q_X1)) | ~term(q_X1) | ~term(X12) [duplicate literal removal 196208]
196296. term(q_Forall(q_X5,q_X1)) | ~term(X12) [subsumption resolution 196249,585]
196622. q_Forall(q_X5,q_X5) = q_Forall(X2,X2) | ~term(X2) | ~term(X2) | ~term(q_X5) [superposition 19681,715]
196765. q_Forall(q_X5,q_X5) = q_Forall(X2,X2) | ~term(X2) | ~term(q_X5) [duplicate literal removal 196622]
196770. q_Forall(q_X5,q_X5) = q_Forall(X2,X2) | ~term(X2) [subsumption resolution 196765,586]
196837. 434 <=> term(q_Forall(q_X5,q_X1)) [avatar definition]
196839. term(q_Forall(q_X5,q_X1)) <- (434) [avatar component clause 196837]
196840. 2 | 434 [avatar split clause 196296,196837,832]
197131. q_Forall(q_X5,q_X1) = q_Forall(q_X5,q_X5) | ~term(q_X1) | ~term(q_X1) [superposition 20247,19766]
197135. q_Forall(q_X5,q_X4) = q_Forall(q_X5,q_X5) | ~term(q_X4) | ~term(q_X4) [superposition 20247,20056]
197136. q_Forall(q_X5,q_X5) = q_Forall(q_c,q_c) | ~term(q_c) | ~term(q_c) [superposition 19670,20247]
197219. q_Forall(q_X5,q_X5) = q_Forall(q_c,q_c) | ~term(q_c) [duplicate literal removal 197136]
197220. q_Forall(q_X5,q_X4) = q_Forall(q_X5,q_X5) | ~term(q_X4) [duplicate literal removal 197135]
197223. q_Forall(q_X5,q_X1) = q_Forall(q_X5,q_X5) | ~term(q_X1) [duplicate literal removal 197131]
197228. q_Forall(q_X5,q_X1) = q_Forall(q_X5,q_X5) [subsumption resolution 197223,585]
197269. ~q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X1))) <- (~320) [backward demodulation 136025,197228]
197283. q_Forall(q_X5,q_X1) = q_Forall(X2,X2) | ~term(X2) [backward demodulation 196770,197228]
197661. q_Forall(q_X5,q_X4) = q_Forall(q_X5,q_X5) [subsumption resolution 197220,587]
197662. q_Forall(q_X5,q_X1) = q_Forall(q_X5,q_X4) [forward demodulation 197661,197228]
197667. q_Forall(q_X5,X111) = sub(q_Forall(q_X5,q_X1),q_X4,X111) | ~term(X111) [backward demodulation 23044,197662]
197882. q_Forall(q_X5,q_X5) = q_Forall(q_c,q_c) [subsumption resolution 197219,583]
197883. q_Forall(q_X5,q_X1) = q_Forall(q_c,q_c) [backward demodulation 197228,197882]
197885. ~q_Truth(q_Forall(q_c,q_c)) <- (~47) [backward demodulation 1740,197883]
197949. q_Forall(q_c,q_c) = q_Forall(X2,q_X1) | ~term(X2) [backward demodulation 195949,197883]
197982. term(q_Forall(q_c,q_c)) <- (434) [backward demodulation 196839,197883]
198009. ~q_Truth(q_Forall(q_X2,q_Forall(q_c,q_c))) <- (~320) [backward demodulation 197269,197883]
198021. q_Forall(q_c,q_c) = q_Forall(X2,X2) | ~term(X2) [backward demodulation 197283,197883]
198176. q_Forall(q_X5,X111) = sub(q_Forall(q_c,q_c),q_X4,X111) | ~term(X111) [backward demodulation 197667,197883]
199459. sub(q_Forall(q_c,q_c),q_X4,X75) = q_Forall(q_c,sub(q_X1,q_X4,X75)) | ~term(q_X1) | ~term(q_c) [superposition 2991,197949]
199648. sub(q_Forall(q_c,q_c),q_X4,X75) = q_Forall(q_c,sub(q_X1,q_X4,X75)) | ~term(q_c) [subsumption resolution 199459,585]
199649. sub(q_Forall(q_c,q_c),q_X4,X75) = q_Forall(q_c,sub(q_X1,q_X4,X75)) [subsumption resolution 199648,583]
199657. q_Forall(q_X5,X111) = q_Forall(q_c,sub(q_X1,q_X4,X111)) | ~term(X111) [backward demodulation 198176,199649]
201694. q_Forall(q_X1,q_c) = q_Forall(q_X1,q_X1) | ~term(q_c) | ~term(q_c) [superposition 20567,5164]
201697. q_Forall(q_X1,q_X1) = q_Forall(q_c,q_c) | ~term(q_c) | ~term(q_c) [superposition 20391,20567]
201779. q_Forall(q_X1,q_X1) = q_Forall(q_c,q_c) | ~term(q_c) [duplicate literal removal 201697]
201780. q_Forall(q_X1,q_c) = q_Forall(q_X1,q_X1) | ~term(q_c) [duplicate literal removal 201694]
201784. q_Forall(q_X1,q_c) = q_Forall(q_X1,q_X1) [subsumption resolution 201780,583]
201898. q_Forall(q_X1,q_X1) = q_Forall(q_c,q_c) [subsumption resolution 201779,583]
201899. q_Forall(q_X1,q_c) = q_Forall(q_c,q_c) [backward demodulation 201784,201898]
213529. q_Forall(q_X5,X0) = q_Forall(q_c,q_X1) | ~term(X0) | ~term(X0) [superposition 199657,625]
213562. q_Forall(q_X5,X0) = q_Forall(q_c,q_X1) | ~term(X0) [duplicate literal removal 213529]
231818. q_Forall(q_c,q_c) = q_Forall(q_c,q_X1) | ~term(q_X5) | ~term(q_X5) [superposition 213562,198021]
231881. q_Truth(q_Forall(q_c,q_X1)) | ~term(q_Neg(q_c)) <- (185) [superposition 11900,213562]
231892. ~q_Truth(q_Forall(q_c,q_X1)) | ~term(q_Neg(q_X5)) <- (~253) [superposition 45076,213562]
232192. q_Forall(q_c,q_c) = q_Forall(q_c,q_X1) | ~term(q_X5) [duplicate literal removal 231818]
232193. q_Forall(q_c,q_c) = q_Forall(q_c,q_X1) [subsumption resolution 232192,586]
238005. q_Truth(q_Forall(q_c,q_c)) | ~term(q_Neg(q_c)) <- (185) [forward demodulation 231881,232193]
238006. ~term(q_Neg(q_c)) <- (~47, 185) [subsumption resolution 238005,197885]
238017. ~q_Truth(q_Forall(q_c,q_c)) | ~term(q_Neg(q_X5)) <- (~253) [forward demodulation 231892,232193]
238541. ~term(q_c) <- (~47, 185) [resolution 238006,622]
238542. $false <- (~47, 185) [subsumption resolution 238541,583]
238543. 47 | ~185 [avatar contradiction clause 238542]
239206. q_Truth(q_Forall(q_c,q_c)) <- (26) [forward demodulation 1125,201899]
239207. $false <- (26, ~47) [subsumption resolution 239206,197885]
239208. ~26 | 47 [avatar contradiction clause 239207]
239894. q_Truth(q_Forall(q_c,q_c)) <- (47) [forward demodulation 1739,197883]
239905. ~q_Truth(q_Forall(q_X5,q_X5)) | ~term(q_Forall(q_X5,q_X5)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) | ~term(q_Quote(sK1(q_Forall(q_X5,q_X5)))) [subsumption resolution 65865,586]
239906. ~q_Truth(q_Forall(q_X5,q_X5)) | ~term(q_Forall(q_X5,q_X5)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) [subsumption resolution 239905,610]
239907. ~q_Truth(q_Forall(q_c,q_c)) | ~term(q_Forall(q_X5,q_X5)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) [forward demodulation 239906,197882]
239908. ~term(q_Forall(q_c,q_c)) | ~q_Truth(q_Forall(q_c,q_c)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) [forward demodulation 239907,197882]
239909. ~q_Truth(q_Forall(q_c,q_c)) | q_Truth(q_Forall(q_X2,q_Forall(q_X5,q_X5))) <- (434) [subsumption resolution 239908,197982]
239910. q_Truth(q_Forall(q_X2,q_Forall(q_c,q_c))) | ~q_Truth(q_Forall(q_c,q_c)) <- (434) [forward demodulation 239909,197882]
239911. ~q_Truth(q_Forall(q_c,q_c)) <- (~320, 434) [subsumption resolution 239910,198009]
239937. $false <- (47, ~320, 434) [subsumption resolution 239894,239911]
239938. ~47 | 320 | ~434 [avatar contradiction clause 239937]
240562. q_Truth(q_Forall(q_c,q_c)) <- (64) [forward demodulation 3064,197882]
242250. ~term(q_Neg(q_X5)) <- (64, ~253) [subsumption resolution 238017,240562]
242609. ~term(q_X5) <- (64, ~253) [resolution 242250,622]
242610. $false <- (64, ~253) [subsumption resolution 242609,586]
242611. ~64 | 253 [avatar contradiction clause 242610]
242617. ~q_Truth(q_Quote(X0)) | ~term(q_Quote(X0)) <- (243) [resolution 40855,653]
242623. ~term(q_Quote(X0)) <- (62, 243) [subsumption resolution 242617,3055]
242624. $false <- (62, 243) [subsumption resolution 242623,610]
242625. ~62 | ~243 [avatar contradiction clause 242624]
242626. $false [avatar sat refutation 834,850,860,870,1094,1135,1167,3066,4622,6446,32001,45077,101178,136026,196840,238543,239208,239938,242611,242625]
% SZS output end Proof for 
% ------------------------------
% Version: Vampire 4.5.1 (commit 57a6f78c on 2020-07-15 11:59:04 +0200)
% Termination reason: Refutation

% Memory used [KB]: 270187
% Time elapsed: 3.524 s
% ------------------------------
% ------------------------------
% Success in time 119.544 s
